# Part 3: Reverse Mode Automatic Differentiation with PyTorch

```python
# Execute this code block to install dependencies when running on colab
try:
    import torch
except:
    from os.path import exists
    from wheel.pep425tags import get_abbr_impl, get_impl_ver, get_abi_tag
    platform = '{}{}-{}'.format(get_abbr_impl(), get_impl_ver(), get_abi_tag())
    cuda_output = !ldconfig -p|grep cudart.so|sed -e 's/.*\.\([0-9]*\)\.\([0-9]*\)$/cu\1\2/'
    accelerator = cuda_output[0] if exists('/dev/nvidia0') else 'cpu'

    !pip install -q http://download.pytorch.org/whl/{accelerator}/torch-1.0.0-{platform}-linux_x86_64.whl torchvision
```

PyTorch implements Dynamic Reverse Mode Automatic Differentiation, much like we did in the previous exercise. There is one really major difference in what PyTorch provides over our simple example: it works directly with matrices (`Tensor`s) rather than with scalars (although obviously a matrix can represent a scalar).

In this tutorial, we'll explore PyTorch's AD implementation. Note that we're using the API of PyTorch 0.4 or later which simplifies use of AD (previous versions required wrapping all `Tensor`s that you wanted to compute gradients of in `Variable` objects; PyTorch 0.4 removes the need to do this and allows `Tensor`s themselves to track gradients).

We'll start with the simple example we tried earlier in the code block below:

__Task:__ Run the following code and verify the solution is correct

```python
import torch

# set up the problem
x = torch.tensor(0.5, requires_grad=True)
y = torch.tensor(4.2, requires_grad=True)
z = x * y + torch.sin(x)

print("z = " + str(z.item()))

z.backward() # this goes through the computation graph and accumulates the gradients in the cached .grad attributes
print("dz/dx = " + str(x.grad.item()))
print("dz/dy = " + str(y.grad.item()))
```

Output:
```
z = 2.57942533493042
dz/dx = 5.077582359313965
dz/dy = 0.5
```

As with our own AD implementation, PyTorch lets us differentiate through an algorithm.

__Task__: Use the block below to compute the gradient $\partial z/\partial x$ of the following pseudocode algorithm and store the result in the `dzdx` variable:

    x = 0.5
    z = 1
    i = 0
    while i<2:
        z = (z + i) * x * x
        i = i + 1

<!-- grade_id: cell-62e1f0838236eddf, points: 0, grade: false, solution: true -->
```python
x = torch.tensor(0.5, requires_grad=True)
z = torch.tensor(1.0, requires_grad=True)

i = 0
while i < 2:
    z = (z + i) * x * x
    i = i + 1
    print("z = " + str(z.item()))

z.backward() 

dzdx = x.grad.item()
dzdx

```

Output:
```
z = 0.25
z = 0.3125
1.5
```

<!-- grade_id: cell-cbe71552690f710b, points: 4, grade: true, solution: false -->
```python
assert dzdx

```

## PyTorch limitations: in-place operations and aliasing

PyTorch will throw an error at runtime if you try to differentiate through an in-place operation on a tensor. 

__Task__: Run the following code to see this in action.

```python
# https://stackoverflow.com/questions/52920098/what-does-the-underscore-suffix-in-pytorch-functions-mean
```

```python
x = torch.tensor(1.0, requires_grad=True)
y = x.tanh()
y.add_(3) # inplace addition
y.backward()
```

Output:
```
~/opt/anaconda3/lib/python3.8/site-packages/torch/tensor.py in backward(self, gradient, retain_graph, create_graph)
    219                 retain_graph=retain_graph,
    220                 create_graph=create_graph)
--> 221         torch.autograd.backward(self, gradient, retain_graph, create_graph)
    222 
    223     def register_hook(self, hook):

~/opt/anaconda3/lib/python3.8/site-packages/torch/autograd/__init__.py in backward(tensors, grad_tensors, retain_graph, create_graph, grad_variables)
    128         retain_graph = create_graph
    129 
--> 130     Variable._execution_engine.run_backward(
    131         tensors, grad_tensors_, retain_graph, create_graph,
    132         allow_unreachable=True)  # allow_unreachable flag

RuntimeError: one of the variables needed for gradient computation has been modified by an inplace operation: [torch.FloatTensor []], which is output 0 of TanhBackward, is at version 1; expected version 0 instead. Hint: enable anomaly detection to find t
```

Aliasing is also something that can't be differentiated through and will result in a slightly more cryptic error.

__Task__: Run the following code to see this in action. If you don't understand what this code does add some `print` statements to show the values of `x` and `y` at various points.

```python
x = torch.tensor([1, 2, 3, 4], requires_grad=True, dtype=torch.float)
y = x[:1]
y.add_(3)
y.backward()

```

Output:
```
~/opt/anaconda3/lib/python3.8/site-packages/torch/tensor.py in backward(self, gradient, retain_graph, create_graph)
    219                 retain_graph=retain_graph,
    220                 create_graph=create_graph)
--> 221         torch.autograd.backward(self, gradient, retain_graph, create_graph)
    222 
    223     def register_hook(self, hook):

~/opt/anaconda3/lib/python3.8/site-packages/torch/autograd/__init__.py in backward(tensors, grad_tensors, retain_graph, create_graph, grad_variables)
    128         retain_graph = create_graph
    129 
--> 130     Variable._execution_engine.run_backward(
    131         tensors, grad_tensors_, retain_graph, create_graph,
    132         allow_unreachable=True)  # allow_unreachable flag

RuntimeError: leaf variable has been moved into the graph interior
```

## Dealing with multiple outputs

PyTorch can deal with the case where there are multiple output variables if we can formulate the expression in terms of tensor operations. Consider the example from the presentation for example:

$$\begin{cases}
     z = 2x + \sin x\\
     v = 4x + \cos x
\end{cases}$$

We could formulate this as:

$$
\begin{bmatrix}z \\ v\end{bmatrix} = \begin{bmatrix}2 \\ 4\end{bmatrix} \odot \bar{x} + \begin{bmatrix}1 \\ 0\end{bmatrix} \odot \sin\bar x + \begin{bmatrix}0 \\ 1\end{bmatrix} \odot \cos\bar x
$$

where 

$$
\bar x = \begin{bmatrix}x \\ x\end{bmatrix}
$$

and $\odot$ represents the Hadamard or element-wise product. This is demonstrated using PyTorch in the following code block.

__Task:__ run the code below.

```python
x = torch.tensor([[1.0],[1.0]], requires_grad=True)

zv = ( torch.tensor([[2.0],[4.0]]) * x +
         torch.tensor([[1.0], [0.0]]) * torch.sin(x) +
         torch.tensor([[0.0], [1.0]]) * torch.cos(x) )
        
zv.backward(torch.tensor([[1.0],[1.0]])) # Note as we have "multiple outputs" we must pass in a tensor of weights of the correct size

print(x.grad)
```

Output:
```
tensor([[2.5403],
        [3.1585]])
```

__Task:__ Use the following box to write down the derivative of the expression for $\begin{bmatrix}z \\ v\end{bmatrix}$ and verify the gradients $\partial z/\partial x$ and $\partial v/\partial x$ are correct for $x=1$.

<!-- grade_id: cell-b85be7c90ef1f408, points: 3, grade: true, solution: true -->
$$
\frac{\partial {z}}{{\partial {x}}} = 2 + cosx\\
\frac{\partial {v}}{{\partial {x}}} = 4 - sinx\\
$$

```python
import math

def dzdx_ana(x):
    return 2 + math.cos(x)

def dvdx_ana(x):
    return 4 - math.sin(x)

dzdx_ana(1), dvdx_ana(1)

```

Output:
```
(2.5403023058681398, 3.1585290151921033)
```

## Gradient descent & gradients with respect to a vector
Let's put everything together and using automatically computed gradients to find the minima of a function by taking steps down the gradient from an initial position. Rather than explicitly creating each input variable as a scalar as in the previous examples, we'll use a vector instead (so our gradients will be with respect to each element of the vector).

__Task:__ work through the following example to see how taking gradients with respect to a vector works & how simple gradient descent optimisation can be implemented.

```python
# This is our starting vector
initial=[[2.0], [1.0], [10.0]]

# This is the function we will optimise (feel free to work out the analytic minima!)
def function(x):
    return x[0]**2 + x[1]**2 + x[2]**2

x = torch.tensor(initial, requires_grad=True, dtype=torch.float)
for i in range(0,2):
    # manually dispose of the gradient (in reality it would be better to detach and zero it to reuse memory)
    x.grad=None
 
    # evaluate the function
    J = function(x) 

    # auto-compute the gradients at the previously evaluated point x
    J.backward()
    print('x:', x.data, x.grad)
    
    # compute the update
    z = x - x.grad*0.1 # new coordinate for the next descent
    print('calculate z: ', x.data, x.grad, z)
    x.data = z
    print('new x:', x.data, x.grad)

    # why must the update in the code above be assigned to a different variable (the z) 
    # before being assigned to the data value of x?
    # x.data = x - x.grad*0.1 # why not this?

#     if i%10 == 0:
#         print((x, function(x).item()))

```

Output:
```
x: tensor([[ 2.],
        [ 1.],
        [10.]]) tensor([[ 4.],
        [ 2.],
        [20.]])
calculate z:  tensor([[ 2.],
        [ 1.],
        [10.]]) tensor([[ 4.],
        [ 2.],
        [20.]]) tensor([[1.6000],
        [0.8000],
        [8.0000]], grad_fn=<SubBackward0>)
new x: tensor([[1.6000],
        [0.8000],
        [8.0000]]) tensor([[ 4.],
        [ 2.],
        [20.]])
x: tensor([[1.6000],
        [0.8000],
        [8.0000]]) tensor([[ 3.2000],
        [ 1.6000],
        [16.0000]])
calculate z:  tensor([[1.6000],
        [0.8000],
        [8.0000]]) tensor([[ 3.2000],
        [ 1.6000],
        [16.0000]]) tensor([[1.2800],
        [0.6400],
        [6.4000]], grad_fn=<SubBackward0>)
new x: tensor([[1.2800],
        [0.6400],
        [6.4000]]) tensor([[ 3.2000],
        [ 1.6000],
        [16.0000]])
```

```python

```

__Task__: Answer the following question in the box below: Why must the update in the code above be assigned to a different variable (the `z`) before being assigned to the `data` value of `x`?

<!-- grade_id: cell-855b4580505f3b5c, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

## Differentiating through random operations

We'll end with an example that will be important later in the course: differentiating with respect to the parameters of a random number generator.

Assume that as some part of a differentiable program that we write we wish to incorporate a random element where we sample values, $z$ from a Normal distribution: $z \sim \mathcal{N}(\mu,\sigma^2)$. We want to learn the parameters of the generator $\mu$ and $\sigma^2$, but how can we do this? In a differentiable program setting we want to differentiate with respect to these parameters, but at first glance it isn't at all obvious what this means as the generator _just_ produces numbers which themselves don't have gradients.

The answer is often called the _reparameterisation trick_: If we note that sampling a Normal distribution is equivalent to drawing numbers from a Uniform distribution and scaling and shifting them: $z \sim \mathcal{N}(\mu,\sigma^2) \equiv z \sim \mu + \sigma\mathcal{U}(0,1)\equiv z = \mu + \sigma \zeta\, \rm{where}\, \zeta\sim\mathcal{U}(0,1)$. With this reparameterisation the gradients with respect to the parameters are obvious.

The following code block demonstrates this in practice; each of the gradients can be interpreted as how much an infintesimal change in $\mu$ or $\sigma$ would change $z$ if we could repeat the sampling operation again with the same value of `torch.rand(1)` being produced:

```python
mu = torch.tensor(1.0, requires_grad=True)
sigma = torch.tensor(1.0, requires_grad=True)

for i in range(10):
    mu.grad = None
    sigma.grad = None
    
    z = mu + sigma * torch.rand(1)
    z.backward()
    print("z:", z.item(), "\tdzdmu:", mu.grad.item(), "\tdzdsigma:", sigma.grad.item())
```

Output:
```
z: 1.4326808452606201 	dzdmu: 1.0 	dzdsigma: 0.4326809048652649
z: 1.4727376699447632 	dzdmu: 1.0 	dzdsigma: 0.4727376699447632
z: 1.2926154136657715 	dzdmu: 1.0 	dzdsigma: 0.2926154136657715
z: 1.5986756086349487 	dzdmu: 1.0 	dzdsigma: 0.5986756086349487
z: 1.2253355979919434 	dzdmu: 1.0 	dzdsigma: 0.22533565759658813
z: 1.1325931549072266 	dzdmu: 1.0 	dzdsigma: 0.13259321451187134
z: 1.696880578994751 	dzdmu: 1.0 	dzdsigma: 0.6968806385993958
z: 1.9604157209396362 	dzdmu: 1.0 	dzdsigma: 0.9604157209396362
z: 1.687875509262085 	dzdmu: 1.0 	dzdsigma: 0.6878754496574402
z: 1.8728947639465332 	dzdmu: 1.0 	dzdsigma: 0.872894823551178
```

```python

```